# Datenstrukturen und Algorithmen

## Praktische Aufgabe 4

In dieser praktischen Aufgabe werden Sie AVL-Bäume implementieren. Nachdem Sie sich in der letzten praktischen Aufgabe mit allgemeinen binären Suchbäumen beschäftigt haben, möchten wir und diesmal mit balancierten Bäumen beschäftigen. Die Hauptaufgabe besteht also darin einen beliebigen binären Suchbaum in einen AVL-Baum umzuwandeln. Sie werden mit Hilfe einiger Teilaufgaben auf die Lösung geführt. 

Die Abgaben werden mit der `nbgrader` Erweiterung korrigiert. Das System erwartet, dass der Code zum Lösen der Aufgaben nach der `#YOUR CODE HERE` Anweisung kommt. Außerdem darf die Zellenreihenfolge nicht geändert werden. Damit Sie selbst Ihre Lösungsvorschläge validieren können, werden Ihnen Unittests zur Verfügung gestellt. Beachten Sie, dass diese Tests keine Garantie sind für das Erreichen der vollen Punktzahl, da Sie nur einen Teil der Funktionalität überprüfen.

Wichtig: Füllen Sie zunächst die erste Zelle mit `#YOUR ANSWER HERE` unter dem Titel `Abgabeteam` mit ihren Namen und Matrikelnummern vollständig aus. Dies ermöglicht uns auch bei technischen Problemen die Abgaben eindeutig zuordnen zu können. Ändern Sie außerdem nicht den Namen der Datei. 

**Übersicht der Aufgaben** (20 Punkte):

1. **AVL-Tree** - insgesamt: 20 Punkte
   - get_balance() - 2P.
   - rotate_left() - 6P.
   - rotate_right() - 6P.
   - balance() - 6P.
 

## Abgabeteam
Bitte füllen Sie die untenstehende Zelle aus mit 

Nummer des Tutoriums,

Voranme Nachname Matrikelnummer 1,

Vorname Nachname Matrikelnummer 2,

(Vorname Nachname Matrikelnummer 3)

<!-- grade_id: cell-7809c128ebe2d0cb, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Module importieren

Zuerst werden die benötigten Module importiert. Sie dürfen keine weiteren Module impotieren.

Wenn in Ihrer Entwickungsumbegung (z.B Google Colab oder Deepnote) bestimmte Module nicht verfügbar sind, dann kommentieren Sie die erste Zeile aus um die Module in der Umgebung zu installieren.

```python
#!pip install nose
from nose.tools import assert_equal
from random import randint
```

# AVL-Tree

Der folgende Code implementiert teilweise einen AVL-Tree.

Ihre Aufgabe ist es die fehlenden Funktionen `get_balance()`, `rotate_left()`, `rotate_right()` und `rebalance()` zu implementieren. 

Der zur Verfügung gestellte Code definiert zunächst die Klass `Node`, sowie die Funktionen `print_tree()`, `get_height()` und `insert()`. Abweichend von der Vorlesung werden wir hier zur Vereinfachung der Implementierung keinen Parent-Pointer nutzen. Zudem speichert jeder Knoten seine Höhe anstatt seines Ungleichgewichts.

```python
class Node():
    def __init__(self, key, left=None, right=None, height=0):
        self.key = key
        self.left = left
        self.right = right
        self.height = height

    def __repr__(self):
        return str(f"({self.left}, {self.key}|{self.height}, {self.right})")


def print_tree(node, level=0):
    """ Generate a very basic graphical output of the tree. """
    if node is not None:
        print_tree(node.right, level + 1)
        print(f"{' ' * 4 * level}->  {str(node.key)}|{node.height}")
        print_tree(node.left, level + 1)


def get_height(node):
    """ Return the height of the tree. """
    if node is None:
        return -1
    return node.height


def insert(key, node=None):
    """ Inserts a node into the given binary search tree and rebalances it afterwards. """
    if node is None:
        return Node(key)
    elif key < node.key:
        node.left = insert(key, node.left)
    elif key > node.key:
        node.right = insert(key, node.right)
    else:  # key == node.key
        pass

    node.height = 1 + max(get_height(node.left), get_height(node.right))
    return rebalance(key, node)

def insert_list(key_list, node=None):
    """ Inserts a list of nodes into the given binary search tree and rebalances it after every insertion. """
    for key in key_list:
        node = insert(key, node)
    return node
```

## a) get_balance() - 2P.

Um die Berechnung des Ungleichgewichts bei jedem Einfügen zu vereinfachen, speichern wir in jedem Knoten die Höhe des Baumes. Dadurch müssen wir jedoch das Ungleichgewicht berechnen können. Implementieren Sie hierzu die Funktion `get_balance()`, die das Ungleichgewicht berechnet. 

Das Ungleichgewicht berechnet sich aus der **Differenz der Höhe des linken und des rechten Teilbaums**. Ist die Höhe des linken Teilbaums also kleiner als die Höhe des rechten, dann ist das Ergebnis negativ, sonst positiv. Sie könnten die Funktion `get_height()` nutzen, um die Höhe eines Teilbaums zu erhalten.

Hinweis: Falls der Knoten, dessen Ungleichgewicht Sie berechnen sollen nicht existiert gibt die Funktion `get_balance()` `0` zurück.

<!-- grade_id: cell-9ca74ce7d933c3ad, points: 0, grade: false, solution: true -->
```python
def get_balance(node):
    """ Return the balance of the tree represented by `node`. """
    if node is None:
        return 0
    else:
        return get_height(node.left) - get_height(node.right)
```

## a) Tests

<!-- grade_id: cell-b37f6ec274c1a240, points: 2, grade: true, solution: false -->
```python
# unittests
assert_equal(get_balance(None), 0)
assert_equal(get_balance(Node(2)), 0)
assert_equal(get_balance(Node(2, left=Node(1), height=1)), 1)
assert_equal(get_balance(Node(2, right=Node(2), height=1)), -1)
assert_equal(get_balance(Node(2, left=Node(1),  right=Node(2), height=1)), 0)
assert_equal(get_balance(Node(3, height=3, left=Node(2, height=2, left=Node(1, height=1, left=Node(0))))), 3)
assert_equal(get_balance(Node(3, height=3, right=Node(2, height=2, right=Node(1, height=1, right=Node(0))))), -3)
```

<!-- grade_id: cell-f89c58c4c099e765, points: 0, grade: true, solution: false -->
```python
# hidden tests

```

<!-- grade_id: cell-cc50b710dc02d153, points: 0, grade: true, solution: false -->
```python
# hidden tests
```

<!-- grade_id: cell-b3b7fa8ba52eaa26, points: 0, grade: true, solution: false -->
```python
# hidden tests
```

<!-- grade_id: cell-8321716771a39482, points: 0, grade: true, solution: false -->
```python
# hidden tests
### BEGIN HIDDEN TEST
#unbalances on the left side
tree_3 = insert_list_bst([5, 4, 3, 2, 1, 0])
assert_equal(get_balance(tree_3), 5)
```

Output:
```
NameError                                 Traceback (most recent call last)
c:\github.com\meesyice\DSALSS22\Abgabe8\blatt08-python.ipynb Cell 15' in <cell line: 4>()
      <a href='vscode-notebook-cell:/c%3A/github.com/meesyice/DSALSS22/Abgabe8/blatt08-python.ipynb#ch0000014?line=0'>1</a> # hidden tests
      <a href='vscode-notebook-cell:/c%3A/github.com/meesyice/DSALSS22/Abgabe8/blatt08-python.ipynb#ch0000014?line=1'>2</a> ### BEGIN HIDDEN TEST
      <a href='vscode-notebook-cell:/c%3A/github.com/meesyice/DSALSS22/Abgabe8/blatt08-python.ipynb#ch0000014?line=2'>3</a> #unbalances on the left side
----> <a href='vscode-notebook-cell:/c%3A/github.com/meesyice/DSALSS22/Abgabe8/blatt08-python.ipynb#ch0000014?line=3'>4</a> tree_3 = insert_list_bst([5, 4, 3, 2, 1, 0])
      <a href='vscode-notebook-cell:/c%3A/github.com/meesyice/DSALSS22/Abgabe8/blatt08-python.ipynb#ch0000014?line=4'>5</a> assert_equal(get_balance(tree_3), 5)

NameError: name 'insert_list_bst' is not defined
```

## b) rotate_left() - 6P.

Um aus einem unbalancierten Baum einen balancierten Baum zu machen, müssen wir rotieren können. Zunächst möchten wir die Rotation nach links implementieren. Schreiben Sie die Funktion `rotate_left()`, die **eine Linksrotation durchführt** und den **neuen Wurzelknoten des Teilbaums zurückgibt**.

Sie können die untenstehende Abbildung als Hilfe nutzen. Der neue Wurzelknoten der zurückgegeben wird ist `y`.

```
Vorher:                     Nachher:
        ->  ...                     
    ->  y                           ->  ...
        ->  z               ->  y 
->  node                            ->  z
    ->  ...                     ->  node
                                    ->  ...
```

Bitte aktualisieren Sie auch die in den Knoten gespeicherte Höhe entsprechend der durchgeführten Rotation. Die Höhe eines Nodes errechnet sich aus **1 + max(Höhe linker Teilbaum, Höhe rechter Teilbaum)**. Hinweis: Überlegen Sie sich zunächst welche Teilbäume ihre Höhe geändert haben und beachten Sie die Abhängigkeiten in der Berechnung der Höhe.

<!-- grade_id: cell-88f6b00cd48fbfad, points: 0, grade: false, solution: true -->
```python
def rotate_left(node):
    """ Perform a rotation to left of the tree based on `node`. """
    root = node.right
    node.right = root.left
    root.left = node
    node.height = 1 + max(get_height(node.left), get_height(node.right))
    root.height = 1 + max(get_height(root.left), get_height(root.right))
    return root
```

## b) Tests

<!-- grade_id: cell-a4639bcbb8e9115b, points: 6, grade: true, solution: false -->
```python
# unittests
for h in [1, 3, 6]:
    tree_1 = Node("n", height=h+3, left=Node(".", height=h), right=Node("y", height=h+2, left=Node("z", height=h), right=Node(".", height=h+1)))
    tree_1 = rotate_left(tree_1)
    assert_equal(str(tree_1), f"(((None, .|{h}, None), n|{h+1}, (None, z|{h}, None)), y|{h+2}, (None, .|{h+1}, None))")
```

## c) rotate_right() - 6P.

Äquivalent möchten wir die Rotation nach rechts implementieren. Implementieren Sie dazu die Funktion `rotate_right()` die **eine Rechtsrotation der Knoten durchführt** und den **neuen Wurzelknoten des Baumes zurückgibt**. Aktualisieren Sie anschließen die Höhe der Teilbäume. Sie können die untenstehende Abbildung als Hilfestellung in der Implementierung nutzen.

```
Vorher:                     Nachher:
                                    ->  ...
        ->  ...                 ->  node
->  node                            ->  z
        ->  z               ->  y
    >  y                        ->  ...
        ->  ...
```

Bitte beachten Sie, die selben Hinweise, wie in der Funktion `rotate_left()`.

<!-- grade_id: cell-8a627f8319bbb6f6, points: 0, grade: false, solution: true -->
```python
def rotate_right(node):
    """ Perform a rotation to right of the tree based on `node`. """
    root = node.left
    node.left = root.right
    root.right = node
    node.height = 1 + max(get_height(node.left), get_height(node.right))
    root.height = 1 + max(get_height(root.left), get_height(root.right))
    return root
```

## c) Tests

<!-- grade_id: cell-8ac327f873d53a63, points: 6, grade: true, solution: false -->
```python
# unittests
for h in [1, 3, 6]:
    tree_1 = Node("n", height=h+3, left=Node("y", height=h+2, left=Node(".", height=h+1), right=Node("z", height=h),), right=Node(".", height=h))
    tree_1 = rotate_right(tree_1)
    assert_equal(str(tree_1), f"((None, .|{h+1}, None), y|{h+2}, ((None, z|{h}, None), n|{h+1}, (None, .|{h}, None)))")
```

## d) rebalance() - 6P.

Implementieren Sie die Funktion `rebalance()`, die nach jedem Aufruf von `insert()` aufgerufen wird, um den neuen binären Suchbaum zu balancieren. Geben Sie am Ende der Implementierung den neuen Wurzelknoten des balancierten Teilbaums zurück. 

Bitte beachten Sie in der Implementierung die vier Fälle die Auftreten können:
- Rotation links
- Rotation rechts
- Doppelte Rotation links / rechts
- Deppelte Rotation rechts / links

Eine Balancierung der Teilbäume ist nur dann notwendig, wenn sich deren Höhe um mehr als eins unterscheidet. Sie können hierzu die zuvor implementierte Funktion `get_balance()` nutzen, um das Ungleichgewicht zu bestimmen.

Falls Sie eine der vorherigen Teilaufgaben nicht lösen konnten, können Sie trotzdem unter der Annahme arbeiten, dass die Funktionen zur Verfügung stehen.

<!-- grade_id: cell-ddd5941afb160b07, points: 0, grade: false, solution: true -->
```python
def rebalance(key, node):
    # YOUR CODE HERE
    raise NotImplementedError()
```

## d) Tests

<!-- grade_id: cell-4584c4bf85adddb5, points: 6, grade: true, solution: false -->
```python
# insert random nodes into a tree and check manually
node_list = [randint(10, 99) for _ in range(15)]
print("Node List:\n", node_list, "\nAVL-Tree:")
print_tree(insert_list(node_list))
```

<!-- grade_id: cell-02d910d6dd9c2f6a, points: 0, grade: true, solution: false -->
```python
#hidden tests
```

<!-- grade_id: cell-e7d1efebe2581ba0, points: 0, grade: true, solution: false -->
```python
#hidden tests
```

<!-- grade_id: cell-f9bdd8da01bc1d61, points: 0, grade: true, solution: false -->
```python
#hidden tests
```

<!-- grade_id: cell-cdb2fde4a211b9e5, points: 0, grade: true, solution: false -->
```python
#hidden tests
```

<!-- grade_id: cell-7288b26c453a223a, points: 0, grade: true, solution: false -->
```python
#hidden tests
```

<!-- grade_id: cell-f64df8dd65bdb40f, points: 0, grade: true, solution: false -->
```python
#hidden tests
```

<!-- grade_id: cell-4cae42054a6e9490, points: 0, grade: true, solution: false -->
```python
#hidden tests
```

## Jupyter Notebook Stolperfalle
Bei der Benutzung von Jupyter Notebooks, wird der globale Zustand aller Variablen zwischen der Ausführung von verschiedenen Zellen erhalten. Dies ist auch der Fall, wenn Zellen gelöscht oder hinzugefügt werden.
Um sicher zu gehen, dass nicht ausversehen notwendige Variablen überschrieben oder gelöscht wurden, kann der Befehl `Kernel -> Restart & Run All` ausgeführt werden.