# Datenstrukturen und Algorithmen

## Praktische Aufgabe 3

In dieser praktischen Aufgabe werden Sie sich mit binären Suchbäumen (BST = binary search tree) beschäftigen. Zunächst implementieren Sie die aus der Vorlesung bekannte Funktion `insert()`, die jetzt neben einem Schlüssel auch einen Wert im Binärbaum (BT = binary tree) abspeichert. Danach überprüfen Sie, ob ein gegebener BT auch ein BST ist und falls nicht, implementieren Sie einen einfachen Algorithmus, um jeden beliebiegen BT in einen BST mit einer vorgegebenen Baumstruktur zu überführen.

Die Abgaben werden mit der `nbgrader` Erweiterung korrigiert. Das System erwartet, dass der Code zum Lösen der Aufgaben nach der `#YOUR CODE HERE` Anweisung kommt. Außerdem darf die Zellenreihenfolge nicht geändert werden. Damit Sie selbst Ihre Lösungsvorschläge validieren können, werden Ihnen Unittests zur Verfügung gestellt. Beachten Sie, dass diese Tests keine Garantie sind für das Erreichen der vollen Punktzahl, da Sie nur einen Teil der Funktionalität überprüfen.

Wichtig: Füllen Sie zunächst die erste Zelle mit `#YOUR ANSWER HERE` unter dem Titel `Abgabeteam` mit ihren Namen und Matrikelnummern vollständig aus. Dies ermöglicht uns auch bei technischen Problemen die Abgaben eindeutig zuordnen zu können. Ändern Sie außerdem nicht den Namen der Datei. 

Zusammenfassung der Aufgaben:

1. **Binäre Suchbäume** - insgesamt: 20 Punkte
   - insert() - 4P.
   - is_bin_search_tree() - 6P.
   - bin_tree_2_list() - 3P.
   - list_2_bin_tree() - 3P.
   - bin_tree_2_bin_search_tree() - 4P.


## Abgabeteam
Bitte füllen Sie die untenstehende Zelle aus mit 

Nummer des Tutoriums,

Voranme Nachname Matrikelnummer 1,

Vorname Nachname Matrikelnummer 2,

(Vorname Nachname Matrikelnummer 3)

<!-- grade_id: cell-28a4063cbff4e8c6, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Module importieren

Als erstes importieren wir die verwendeten Module. Sie dürfen keine weiteren Python Module importieren oder verwenden, als die hier spezifizierten Module.

Falls Sie eine Entwicklungsumgebung, wie Google Colab oder Deepnote verwendeen und bestimmte Module nicht verfügbar sind, kommentieren Sie die erste Zeile aus, um die Module in der Laufzeitumgebung temporär zu installieren.

```python
#!pip install nose

from nose.tools import assert_equal
```

# Binäre Suchbäume

Im folgenden werden Sie einen BST implementieren, der sowohl ein Attribut `key`, als auch `value` besitzt. Dabei wird die Sortierung der Blätter über den Schlüssel `key` bestimmt, während das eigentliche Element in `value` gespeichert wird. Dadurch können wir verschiedene Elemente nach einem beliebigen Schlüssel im Baum speichern. Der unten stehende Code definiert die Struktur von einem Node. Die Funktion `print_tree` gibt jeweils den BT als einen formattierten Output aus.
Bitte beachten Sie, dass wir in dieser Implementierung keinen Parent-Pointer nutzen. Alle Aufgaben lassen sich auch ohne diese Referenz effizient lösen.

```python
class Node:
    def __init__(self, key, value=None, left=None, right=None):
        self.key = key
        self.value = value
        self.left = left
        self.right = right
        

    def __repr__(self):
        return f"({self.left}, {self.key}|{self.value}, {self.right})"


def print_tree(node, level=0):
    if node is not None:
        print_tree(node.right, level + 1)
        print(f"{' ' * 8 * level}->  {str(node.key)}|{str(node.value)}")
        print_tree(node.left, level + 1)
```

```python
print("Valid BST:")
valid_bst = Node(4, 421, Node(2, 123, Node(1, "asd"), Node(3, None, right=Node(3, "vwa"))), Node(6, 234, Node(5, "abc"), Node(7, "abv", right=Node(8, "cde"))))
print_tree(valid_bst)
print(valid_bst)

print("\n------------------\n")

print("Invalid BST:")
invalid_bst = Node(1, 421, Node(2, "kve", Node(4, "abc", Node(0, "esi"))), Node(9, "agg", right=Node(4, "dnf", Node(2, 123))))
print_tree(invalid_bst)
print(invalid_bst)
```

Output:
```
Valid BST:
                        ->  8|cde
                ->  7|abv
        ->  6|234
                ->  5|abc
->  4|421
                        ->  3|vwa
                ->  3|None
        ->  2|123
                ->  1|asd
(((None, 1|asd, None), 2|123, (None, 3|None, (None, 3|vwa, None))), 4|421, ((None, 5|abc, None), 6|234, (None, 7|abv, (None, 8|cde, None))))

------------------

Invalid BST:
                ->  4|dnf
                        ->  2|123
        ->  9|agg
->  1|421
        ->  2|kve
                ->  4|abc
                        ->  0|esi
((((None, 0|esi, None), 4|abc, None), 2|kve, None), 1|421, (None, 9|agg, ((None, 2|123, None), 4|dnf, None)))
```

## a) insert() - 4P.

Implementieren Sie die Funktion `insert()`, die `node`, `key`, sowie `value` als Parameter übergegeben bekommt. `node` ist hierbei der Elternknoten in den ein Element `value` entsprechend des Schlüssels `key` eingefügt wird. Der resultierende Baum soll dabei die Eigenschaften eines BST behalten. Orientieren sie sich an der aus der Vorlesung bekannten Funktion `insert()` und wandeln diese an den nötigen Stellen ab.

Wenn zweimal der gleiche `key` eingefügt wird, dann wird einfach nur `value` ersetzt. 

<!-- grade_id: cell-d4538ec7909a57aa, points: 0, grade: false, solution: true -->
```python
def insert(node, key, value):
    """ Insert a new node into a binary search tree based on `key` which holds `value`. """
    # YOUR CODE HERE
    
    newNode = Node(key, value)

    if not node:
        node = Node(key, value)
    elif key > node.key and node.right:
           insert(node.right, key, value)
    elif key < node.key and node.left:
           insert(node.left, key, value)
    elif key > node.key and not node.right:
        node.right = newNode
    elif key < node.key and not node.left:
        node.left = newNode

    return node
    #raise NotImplementedError()
```

## a) Tests

<!-- grade_id: cell-5985aec0890b8f05, points: 4, grade: true, solution: false -->
```python
root = None 
root = insert(root, 3, 1177)
assert_equal(str(root), "(None, 3|1177, None)")

```

<!-- grade_id: cell-45f5d7f00d5bc60e, points: 0, grade: true, solution: false -->
```python
root = None 
root = insert(root, 0, "Value")
root = insert(root, 1, "Value")
root = insert(root, 2, "Value")
assert_equal(str(root), "(None, 0|Value, (None, 1|Value, (None, 2|Value, None)))")

```

<!-- grade_id: cell-b467a9a56c876490, points: 0, grade: true, solution: false -->
```python
# hidden tests
```

Die nächste Aufgabe besteht darin einen beliebigen BT in einen BST umzuwandeln. Dabei darf die Form des Baumes, also die Verzweigung, nicht verändert werden.

Hierzu werden Sie zunächst eine Funktion implementieren, die überprüft ob ein gegebener BT auch ein BST ist. Danach implementieren Sie zwei Funktionen, die eine Inorder-Traversal über den BT durchführen und die Nodes in eine Liste schreiben bzw. Nodes aus einer Liste in einen BT überführen. Anschließend nutzen Sie diese Funktionen, sowie das Sortieren von Listen, um einen beliebigen BT in einen BST umzuwandeln und dabei die Verzweigungen des BT nicht zu verändern.

Der zu implementierende Algorithmus besteht aus den folgenden Schritten:
- Falls der BT ein BST ist:
  - mache nichts und geben den BT zurück
- Sonst, führe die folgenden Schritte aus:
  - Konvertiere den BT mit Inorder-Traversal in eine Liste
  - Sortiere die Liste nach dem Schlüssel
  - Konvertiere die Liste mit Inoder-Traversal in einen BT

## b) is_bin_search_tree() - 6P.

Implementieren Sie die Funktion `is_bin_search_tree()`, die `node` als Parameter übergegeben bekommt und überprüft, ob es sich um den BT der durch `node` dargestellt wird, um einen BST handelt. Die Funktion soll `True` zurückgeben, falls ein BST übergegeben wurde, sonst gibt sie `False` aus.

```
Eingabe:
Node(4, None, Node(2), Node(6))

Ausgabe:
True


Eingabe:
Node(3, None, Node(2,None, Node(1), Node(4)), Node(5))

Ausgabe:
False
```

<!-- grade_id: cell-4473dc82bd7f04f7, points: 0, grade: false, solution: true -->
```python
def is_bin_search_tree(node):
    """ Check if a given binary tree is a binary search tree. """
    # YOUR CODE HERE
    if node.right and node.left:
        isbin = node and node.right.key > node.key and node.left.key < node.key
        return isbin and is_bin_search_tree(node.left) and is_bin_search_tree(node.right)
    else:
        return node is not None
    #raise NotImplementedError()
    
```

## b) Tests

<!-- grade_id: cell-fa6aa67562fbf344, points: 6, grade: true, solution: false -->
```python
# unittests
valid_bst = Node(4, 421, Node(2, 123, Node(1, "asd"), Node(3)), Node(6, 234, Node(5, "abc"), Node(7, "abv", right=Node(8, "cde"))))
assert_equal(True, is_bin_search_tree(valid_bst))

invalid_bst = Node(1, 421, Node(2, "kve", Node(4, "abc", Node(0, "esi"))), Node(9, "agg", right=Node(4, "dnf", Node(2, 123))))
assert_equal(False, is_bin_search_tree(invalid_bst))
```

<!-- grade_id: cell-7c376d9043defa12, points: 0, grade: true, solution: false -->
```python
# hidden tests
```

<!-- grade_id: cell-9a00280a3e3bdc4a, points: 0, grade: true, solution: false -->
```python
# hidden tests
```

<!-- grade_id: cell-bb1b714fbd171b6f, points: 0, grade: true, solution: false -->
```python
# hidden tests
```

## c) bin_tree_2_list() - 3P.

Implementieren Sie die Funktion `bin_tree_2_list()`, die einen BT `node` erhält und eine Liste zurückgibt. Diese Liste enthält Tupel von Schlüsseln und Werten als Elemente, die jeweils die Nodes des BT in Inorder-Traversal beschreiben.

```
Eingabe:
Node(4, 421, Node(2, 123, Node(1, "asd"), Node(3, None, right=Node(3, "vwa"))), Node(6, 234, Node(5, "abc"), Node(7, "abv", right=Node(8, "cde"))))

Ausgabe:
[(1, 'asd'), (2, 123), (3, None), (3, 'vwa'), (4, 421), (5, 'abc'), (6, 234), (7, 'abv'), (8, 'cde')]


Eingabe:
Node(1, 421, Node(2, "kve", Node(4, "abc", Node(0, "esi"))), Node(9, "agg", right=Node(4, "dnf", Node(2, 123))))

Ausgabe:
[(0, 'esi'), (4, 'abc'), (2, 'kve'), (1, 421), (9, 'agg'), (2, 123), (4, 'dnf')]
```

<!-- grade_id: cell-d24b7fb71a44e22a, points: 0, grade: false, solution: true -->
```python
def bin_tree_2_list(node):
        """ Convert keys & values of a binary tree into a list. """
        # YOUR CODE HERE
        raise NotImplementedError()
```

## c) Tests

<!-- grade_id: cell-94757a37753d1176, points: 3, grade: true, solution: false -->
```python
valid_bst = Node(4, 421, Node(2, 123, Node(1, "asd"), Node(3, None, right=Node(3, "vwa"))), Node(6, 234, Node(5, "abc"), Node(7, "abv", right=Node(8, "cde"))))
assert_equal([(1, 'asd'), (2, 123), (3, None), (3, 'vwa'), (4, 421), (5, 'abc'), (6, 234), (7, 'abv'), (8, 'cde')],
             bin_tree_2_list(valid_bst))


invalid_bst = Node(1, 421, Node(2, "kve", Node(4, "abc", Node(0, "esi"))), Node(9, "agg", right=Node(4, "dnf", Node(2, 123))))
assert_equal([(0, 'esi'), (4, 'abc'), (2, 'kve'), (1, 421), (9, 'agg'), (2, 123), (4, 'dnf')],
             bin_tree_2_list(invalid_bst))
```

<!-- grade_id: cell-c763e2a663d7a14c, points: 0, grade: true, solution: false -->
```python
# hidden tests
```

## d) list_2_bin_tree() - 3P.

Implementiere Sie die Funktion `list_2_bin_tree()`, die als Parameter einen BT `node` und eine Liste `node_list` übergegeben bekommt. 
Hierbei werden die Tupel aus Schlüssel und Wert aus der Liste `node_list` mit Inorder-Traversal in den BT eingesetzt.

<!-- grade_id: cell-7d893db96c304554, points: 0, grade: false, solution: true -->
```python
def list_2_bin_tree(node, node_list):
    """ Convert keys & values from a list into a binary tree. """
    # YOUR CODE HERE
    raise NotImplementedError()
```

## d) Tests

<!-- grade_id: cell-a3a058f4a5f53af0, points: 3, grade: true, solution: false -->
```python
tree = Node(0, 0, Node(1, 1, Node(2, 2), Node(3, 3)), Node(4, 4, Node(5, 5), Node(6, 6, Node(7, 7))))
list = [(1, 'asd'), (2, 123), (3, None), (3, 'vwa'), (4, 421), (5, 'abc'), (6, 234), (7, 'abv'), (8, 'cde')]
list_2_bin_tree(tree, list)
assert_equal(str(tree), "(((None, 1|asd, None), 2|123, (None, 3|None, None)), 3|vwa, ((None, 4|421, None), 5|abc, ((None, 6|234, None), 7|abv, None)))")

tree = Node(0, 0, Node(1, 1, Node(2, 2, Node(3, 3, Node(4, 4, Node(5, 5, Node(6, 6)))))))
list = [(0, 'esi'), (4, 'abc'), (2, 'kve'), (1, 421), (9, 'agg'), (2, 123), (4, 'dnf')]
list_2_bin_tree(tree, list)
assert_equal(str(tree), "(((((((None, 0|esi, None), 4|abc, None), 2|kve, None), 1|421, None), 9|agg, None), 2|123, None), 4|dnf, None)")
```

## e) bin_tree_2_bin_search_tree() - 4P.

Implementieren Sie die Funktion `bin_tree_2_bin_search_tree()`, die überprüft ob ein übergebener BT `node` auch ein BST ist und falls dies nicht der Fall ist, diesen in einen BST überführt. Hierbei soll die ursprüngliche Form (Verzweigungen) des BT erhalten bleiben. Nutzen Sie bei der Lösung dieser Aufgabe, die von Ihnen zuvor implementierten Funktionen. Sie können davon ausgehen, dass wir als Eingabe nur valide BT verwenden, die z.B. keine doppelten Nodes besitzen.

Hinweis: Da der Fokus dieser Aufgabe nicht auf den Sortieralgorithmen selbst liegt, dürfen Sie Funktionen wie `sort()` oder `sorted()` in Ihrer Lösung verwenden.

<!-- grade_id: cell-8e1c8cb84c936097, points: 0, grade: false, solution: true -->
```python
def bin_tree_2_bin_search_tree(node):
    """ Convert a binary tree into a binary search tree. Preserves the original structure of the binary tree. """
    # YOUR CODE HERE
    raise NotImplementedError()
```

## e) Tests

<!-- grade_id: cell-9d83901cf4e7e5e7, points: 4, grade: true, solution: false -->
```python
invalid_bst = Node(1, 421, Node(2, "kve", Node(4, "abc", Node(0, "esi"))), Node(9, "agg", right=Node(7, "dnf", Node(3, 123))))
bin_tree_2_bin_search_tree(invalid_bst)
assert_equal(True, is_bin_search_tree((invalid_bst)))
```

<!-- grade_id: cell-5208d65b8cc53aa3, points: 0, grade: true, solution: false -->
```python
# hidden tests
```

## Jupyter Notebook Stolperfalle
Bei der Benutzung von Jupyter Notebooks, wird der globale Zustand aller Variablen zwischen der Ausführung von verschiedenen Zellen erhalten. Dies ist auch der Fall, wenn Zellen gelöscht oder hinzugefügt werden.
Um sicher zu gehen, dass nicht ausversehen notwendige Variablen überschrieben oder gelöscht wurden, kann der Befehl `Kernel -> Restart & Run All` ausgeführt werden.